Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:<br/>
**Do not change the name of the file and upload in folder "c" on Google Drive**

```python
NAME = "Neha"
```

---

# Assignment 3

Create a class `Planet` which takes two variables while constructing an object i.e. `radius (in meters)` and `rotation_period (in seconds)` <br>
Now implement 4 methods
```
Q1. surface_area
Q2. rotation_frequency
Q3. A function such that it checks for equality between two planets (i.e. planet1 == planet2)
```
\begin{equation*}
{rotation\_frequency}=\frac{2{\pi}}{rotation\_period}
\end{equation*}
<br>For 4th read <a href="https://docs.python.org/3/reference/datamodel.html">This</a>

<!-- grade_id: cell-6d5c9250b66244b0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE BELOW
class Planet:
    
    pi = 3.141592653589793
    
    def __init__ (self, radius, rotation_period):
        self.radius = radius
        self.r_period = rotation_period
        
    def surface_area(self):
        return 4*self.pi*(self.radius**2)
    
    def rotation_frequency(self):
        return 2*self.pi/self.r_period
    
    def __eq__ (self, other_planet):
        if(self.radius == other_planet.radius  and  self.r_period == other_planet.r_period):
            return True
        else:
            return False
```

<!-- grade_id: cell-561ce6b19c552aff, points: 1, grade: true, solution: false -->
```python
#This cell should pass with no errors
earth = Planet(6.3781e6, 86400)
jupyter = Planet(69.911e6, 9.925*60*60)
import pytest
assert earth.surface_area() == pytest.approx(5.112015e14)
assert jupyter.surface_area() == pytest.approx(6.14187e16)
```

<!-- grade_id: cell-1e483767c073a280, points: 1, grade: true, solution: false -->
```python
assert earth.rotation_frequency() == pytest.approx(7.2722e-5)
assert jupyter.rotation_frequency() == pytest.approx(1.758518e-4)
```

<!-- grade_id: cell-d872b8d55154e343, points: 1, grade: true, solution: false -->
```python
earth2 = Planet(6.3781e6, 86400)
assert earth == earth2
```

## Q4
Now a catastrophe happened and the value of &pi; got changed to 4.5 <br>
           Find the new surface area. <br>
           **This catastrophe must effect every planet.**

<!-- grade_id: cell-78a0c5fcb8f34ab9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE BELOW
Planet.pi = 4.5
```

<!-- grade_id: cell-5354b684fdf49890, points: 1, grade: true, solution: false -->
```python
assert earth.surface_area() == 732242872980000.0
assert jupyter.surface_area() == 87975862578000000
```

## Q5
Return a 3x3 matrix with values ranging from *start* to *stop* both inclusive<br>
**Note: stop-start would always be 8**<br>
Complete `f1()`

<!-- grade_id: cell-4bbfc9749c7fc9d7, points: 0, grade: false, solution: true -->
```python
import numpy as np
def f1(start,stop):
    # YOUR CODE BELOW
    f = np.arange(start,stop+1)
    f = f.reshape(3,3)
    return f
    pass
```

<!-- grade_id: cell-826ec4211109b15c, points: 1, grade: true, solution: false -->
```python
#Tests
import numpy as np
assert np.array_equal(f1(2,10),np.array([[2,3,4],[5,6,7],[8,9,10]]))
```

## Q6
A 2D array `m` is given to the function, return the last column only.
Complete `f2()`

<!-- grade_id: cell-c9c4d6b2904b8d78, points: 0, grade: false, solution: true -->
```python
def f2(m):
    # YOUR CODE BELOW
    dim = np.shape(m)
    ncol=dim[1]-1
    return m[:,ncol]

```

<!-- grade_id: cell-ad52563149a6d0ad, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(f2(np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])),np.array([3, 6, 9]))
assert np.array_equal(f2(np.array([[11, 21, 31], [84, 75, 56], [97, 98, 99]])),np.array([31, 56, 99]))
```

```python

```