Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:<br/>
**Do not change the name of the file and upload in folder "b" on Google Drive**

```python
NAME = "Neha"
```

---

# Assignment 2

## Q1
Given a number *n*, and a function ```f1()```, you must complete the function to `return` a list from [1,n]

<!-- grade_id: cell-6d5c9250b66244b0, points: 0, grade: false, solution: true -->
```python
def f1(n):
    # YOUR CODE BELOW
    myList = []
    for i in range(1,n+1):
        myList.append(i)
    return myList
    pass
```

<!-- grade_id: cell-561ce6b19c552aff, points: 1, grade: true, solution: false -->
```python
#Tests
assert f1(2)==[1,2]
assert f1(3)==[1,2,3]
```

## Q2
Write a function f2(), which checks whether a number n is prime or not, it must return ```True``` if prime, else ```False``` if not a prime. <br>Note: 1 is considered to be Non-Prime

<!-- grade_id: cell-4bbfc9749c7fc9d7, points: 0, grade: false, solution: true -->
```python
def f2(n):
    # YOUR CODE BELOW
    if(n==1):
        return False
    if(n==2):
        return True
    m=n//2
    for i in range(2,m+1):
        if(n%i==0):
            return False
        else:
            return True
    pass
```

<!-- grade_id: cell-826ec4211109b15c, points: 1, grade: true, solution: false -->
```python
#Tests
assert f2(2)==True
assert f2(4)==False
```

## Q3
HackTheBox never gets tired of encoding and decoding strings, here is another question, this one is a little straight forward <br><br> *The string was splitted into a list such that every char was an element in list, then some random chars were added in the beginning and in the end such that the original character stayed in the middle* <br><br>
You are given a function *f3*() and the list *l*, `return` the original string

<br><code>Ex1: if the list is ['abc','pkonh','aqwetmkoi'], then the middle char of each element is 'b','o','t' repectively hence the original string is "bot"</code>

<!-- grade_id: cell-61aba2b1b33f8ef7, points: 0, grade: false, solution: true -->
```python
def f3(l):
    # YOUR CODE BELOW
    s=''
    for i in l:
        s+=i[(len(i)//2)]
    return s
    pass
```

<!-- grade_id: cell-6240053ada5a86d4, points: 1, grade: true, solution: false -->
```python
#Tests
assert f3(['abc','pkonh','aqwetmkoi'])=="bot"
assert f3(['abc','pkonh','o','pkp'])=="book"
```

## Q4
This is an interesting question, you are given a string `s` all in lower case<br>
Make a dictionary storing the frequency of each letter in the string<br>
Complete `f4()` and **return** the dictionary<br>
*Eg:*<br>
*If s="abc"<br>
d={'a':1,'b':2,'c':3}*
<br>And so on

<!-- grade_id: cell-847c586267ef5a2b, points: 0, grade: false, solution: true -->
```python
def f4(s):
    # YOUR CODE BELOW
    l=list(s)
    sl=set(l)
    d={}
    
    for letter in sorted(sl):
        d[letter]=l.count(letter)
    return d
    pass
```

<!-- grade_id: cell-a2473c415456b7d7, points: 1, grade: true, solution: false -->
```python
#Tests
assert f4('abc')=={'a': 1, 'b': 1, 'c': 1}
assert f4('aabcbc')=={'a': 2, 'b': 2, 'c': 2}
```

## Q5 (Optional)
Now, let us come to the real question <br>

HackTheBox really loves strings a lot but they likes palindromic strings more. Today, they took two strings <b>X</b> and <b>Y</b>, each consisting of lower case alphabets.

Now they took a substring from X which is not empty let us say `x` and they took another substring from Y (which is again non empty) which is let us say `y` and then they did the concatenation of `x and y` and called it `z`
Now they want to know whether this `z` is pallindrome or not <br><br>
You will be given two string X and Y and return ```True``` if it possible to choose such string <b>z</b> else ```False```
<div style="text-align: right">-- 3 points </div>

<!-- grade_id: cell-7b8c5a1796c5b1fe, points: 0, grade: false, solution: true -->
```python
def f5(A,B): #A and B are strings
    # You may use f4()
    # YOUR CODE BELOW
    letterCount = f4(A+B)
    count_odd=0
    
    for i in letterCount.values():
        if i%2!=0:
            count_odd+=1
            
    if count_odd <=1:
        return True
    else:
        return False
    pass
```

<!-- grade_id: cell-33cb3c0010494a5b, points: 3, grade: true, solution: false -->
```python
#Tests
assert f5('abc','abc')==True
assert f5('a','b')==False
assert f5('abba','baab')==True
```

## Q6
You are given two integers a and b.

In one move, you can choose some integer *k* from *1 to 10* and <b>add it to a or subtract it from a</b>. <br><br>In other words, you choose an integer k∈[1:10] and perform ```a=a + k``` or ```a = a − k```. You may use different values of k in different moves.

Your task is to **return** the minimum number of moves required to obtain <em><b>b</b></em> from <em><b>a</b></em>.<br>

<div style="font-size: 85%;">PS: You may not need loops, may not</div>

<!-- grade_id: cell-c7e374efc2711302, points: 0, grade: false, solution: true -->
```python
def f6(a,b):
    # YOUR CODE BELOW
    c=abs(b-a)
    if c%10 == 0:
        return c//10
    else:
        return c//10 + 1
    pass
```

<!-- grade_id: cell-24536de406e7c539, points: 1, grade: true, solution: false -->
```python
# Tests
assert f6(5,5)==0    #nothing required
assert f6(13,42)==3  #13→23→32→42  (add 10, add 9, add 10).
assert f6(18,4)==2   #18→10→4  (subtract 8, subtract 6)
```

## Q7
HackTheBox can not get enough of strings. They took a string s and did this and gave you `encoded`
```
l=list(s)
for i in range(len(l)):
    l[i]=chr(ord(l[i])+i)
encoded="".join(l)
```
<br>Can you **return** the original string through `f7(encoded)`?<br>
Read about ord() <a href="https://docs.python.org/3.4/library/functions.html?highlight=ord#ord">From here</a><br>
Read about chr() <a href="https://docs.python.org/3.4/library/functions.html?highlight=ord#chr">From here</a>

<!-- grade_id: cell-13867e36ce8ccbd0, points: 0, grade: false, solution: true -->
```python
def f7(encoded):
    # YOUR CODE BELOW
    l=list(encoded)
    for i in range(len(l)):
        l[i]=chr(ord(l[i])-i)
    s="".join(l)
    return s
    pass
```

<!-- grade_id: cell-1d171295faece88a, points: 1, grade: true, solution: false -->
```python
# Tests
assert f7("def")=="ddd"   
assert f7("sqcfittv~j")=="spaceonova"
assert f7("Pzvkss")=="Python"
```

```python

```